#### Bi 410/510 (Fall 2019)

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

**Group Members**

If this is a group submission edit this cell to add the names and e-mail addresses of the othrer people who worked on the project.
* name (xxx@uoregon.edu)
* name (xxx@uoregon.edu)
* name (xxx@uoregon.edu)

```python

```

## <span style="color:steelblue;">Milestone 2a

### <span style="color:steelblue;">ERA</span>

In baseball a pitcher’s earned run average (ERA) is defined to be the number of earned runs scored by the other team for every nine innings pitched. The formula is

$$ \mathrm{ERA} = 27 \times \frac{\mathrm{runs}}{\mathrm{outs}} $$

For example, if a pitcher is in the game for 5 full innings (which is 15 outs) and the other team scores 2 runs, the ERA would be 3.6 (projecting he would allow an average of 3.6 runs if he pitched an entire 9 innings). 

Fill in the body of this definition of a function that computes ERA.  Arguments passed to the function will be the number of runs and outs, and the result returned should be the ERA rounded to two places.

<!-- grade_id: era, points: 3, grade: true, solution: true -->
```python
def era(runs, outs):
    # YOUR CODE HERE
    ERA= 27*(runs/outs)
    ERA = float("%0.2f" % (ERA))
    return ERA
```

##### Tests 

Pitch 5 full innings (15 outs), give up 2 runs, ERA should be 3.6.

<!-- grade_id: cell-9305034026ed1261, points: 1, grade: true, solution: false -->
```python
assert era(2,15) == 3.6
```

Pitch 6 1/3 innings (19 outs), give up 4 runs, ERA should be 5.68.

<!-- grade_id: cell-116a4c44d874e69f, points: 1, grade: true, solution: false -->
```python
assert era(4,19) == 5.68
```

### <span style="color:steelblue;">Rolling Dice</span>

Python's `random` library has a function named `randint`.  When we pass `randint` two integers $x$ and $y$ it will return a random integer $x <= i <= y$.

We can use `randint` as part of a program that plays a game like Monopoly where players roll two dice.  Fill in the definition of a function named `roll` so it calls `randint` twice, each time getting a random number between 1 and 6, and returns the sum of the two calls.

**Optional:** your function can include a `print` statement that shows the results of the two calls to `randint`, _e.g._
```
>>> roll()
6 + 3 = 9
```

<!-- grade_id: roll, points: 3, grade: true, solution: true -->
```python
from random import randint
def roll():
    # YOUR CODE HERE
    x= randint(1,6)
    y= randint(1,6)
    result= x+y
    print(x ," + ", y, "= ",result)
    return(result)
```

##### Tests 

Each time you call `roll` you should get back a number between 2 and 12.

<!-- grade_id: cell-e58641aa17c521ec, points: 1, grade: true, solution: false -->
```python
assert 2 <= roll() <= 12
```

Output:
```
6  +  4 =  10
```

<!-- grade_id: cell-f5cd45d48b25fafb, points: 1, grade: true, solution: false -->
```python
assert 2 <= roll() <= 12
```

Output:
```
6  +  2 =  8
```

### <span style="color:steelblue;">ORF</span>

An _open reading frame_ (ORF) is a DNA sequence that begins with the start codon `ATG` and ends with one of the stop codons `TAA`, `TAG`, or `TGA` (another requirement is that the sequence length is a multiple of three, but we'll save that for a future project).

Use the code cell below to write the definition of a function named `orf` that will be passed a single argument named `seq`.  The function should return True if `seq` is an ORF and False if it is not an ORF.

<!-- grade_id: orf, points: 3, grade: true, solution: true -->
```python
def orf(seq):
    # YOUR CODE HERE
    if(seq[:3]=="ATG" and(seq[-3:]=="TAA" or seq[-3:]=="TAG" or seq[-3:]=="TGA")):
#         if(seq[-3:]=="TAA" or seq[-3:]=="TAG" or seq[-3:]=="TGA"):
        return True
    return False
```

##### Tests 

This string is an ORF sequence:

<!-- grade_id: cell-dc7c98e0a12e4694, points: 1, grade: true, solution: false -->
```python
assert orf('ATGCAATGGGGATAA')
```

None of these strings are ORFs:

<!-- grade_id: cell-7c9d42c2080534a8, points: 1, grade: true, solution: false -->
```python
assert not orf('GATTACA')
```

<!-- grade_id: cell-51f5cc75bea52a7a, points: 1, grade: true, solution: false -->
```python
assert not orf('ATGAAATTTGGG')
```

<!-- grade_id: cell-40ac0ad2b5449722, points: 1, grade: true, solution: false -->
```python
assert not orf('CCCTAG')
```

### <span style="color:red;">Short Answer Question</span>

Don't forget to complete the short-answer question on the exam page.

Double-click on the markdown cell below and enter your answer in this cell.

<!-- grade_id: cell-e6aff85b4ff2c591, points: 3, grade: true, solution: true -->
1- The repeat function returns 'no no no'

2- The repeat function returns 'yeah'. An error does not occur since the function always returns a string s at least once. s*0 means that the string will be printed zero times but +s part causes the input string to be printed at least once.

```python

```